// intro2.rs
//
// Make the code print a greeting to the world.
//
// Execute `rustlings hint intro2` or use the `hint` watch subcommand for a
// hint.

// I AM NOT DONE

fn main() {
    printline!("Hello there!")
}

//Compiler hint tells us that  "printline macro was not defined in the current scope, but there is a similar macro named println is available".

//So lets try to use use println! then ! 
